**Q1:** Create a Rectangle class that can calculate the area and perimeter of the shape to ensure the program passes the tests below.

<!-- grade_id: cell-983997ad6af236c9, points: 0, grade: false, solution: true -->
```python
class Rectangle:
    def __init__(self, length, width):
        self.length = length
        self.width = width
    
    def calculate_area(self):
        return self.length * self.width
    
    def calculate_perimeter(self):
        return 2 * (self.length + self.width)
```

```python
assert Rectangle(2,3).length == 2
assert Rectangle(2,3).width == 3
assert Rectangle(2,3).calculate_area() == 6
assert Rectangle(3,4).calculate_area() == 12
assert Rectangle(2,3).calculate_perimeter() == 10
assert Rectangle(3,4).calculate_perimeter() == 14
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-239c3cd31216c376, points: 0, grade: true, solution: false -->
```python
assert Rectangle(2,3).length == 2
assert Rectangle(2,3).width == 3
assert Rectangle(2,3).calculate_area() == 6
assert Rectangle(3,4).calculate_area() == 12
assert Rectangle(2,3).calculate_perimeter() == 10
assert Rectangle(3,4).calculate_perimeter() == 14
```

**Q2:** Create the following methods in the provided section of the ShoppingCart class to ensure the program works correctly according to various test cases:

- `add_item(item_name: str, qty: int)`: Adds a specified quantity (qty) of an item (item_name) to the shopping cart.
- `remove_item(item_name: str, qty: int)`: Removes a specified quantity (qty) of an item (item_name) from the shopping cart. If the quantity of an item becomes 0 or negative, remove the item from the cart.
- `get_total_qty()`: Calculates the total quantity of items in the shopping cart.

<!-- grade_id: cell-743ad1759a25cd22, points: 0, grade: false, solution: true -->
```python
class ShoppingCart:
    def __init__(self):
        self.cart = {}
        
    def add_item(self, item_name, qty):
        if item_name not in self.cart:
            self.cart[item_name] = qty
        else:
            self.cart[item_name] += qty
    
    def remove_item(self, item_name, qty):
        try:
            if self.cart[item_name] <= qty:
                self.cart.pop(item_name)
            else:
                self.cart[item_name] -= qty
        except:
            print(f'There is no {item_name} in the cart!!!')
    
    def get_total_qty(self):
        total_qty = 0
        for value in self.cart.values():
            total_qty += value
        return total_qty
```

<!-- grade_id: cell-b1626d544a63028d, points: 0, grade: true, solution: false -->
```python
# Test cases
mycart = ShoppingCart()

mycart.add_item("Papaya", 100)
mycart.add_item("Guava", 200)
mycart.add_item("Orange", 150)

assert len(mycart.cart) == 3
assert mycart.cart['Papaya'] == 100
assert mycart.cart['Guava'] == 200
assert mycart.cart['Orange'] == 150

assert mycart.get_total_qty() == 450
```

```python
# Test cases
mycart = ShoppingCart()

mycart.add_item("Papaya", 100)
mycart.add_item("Guava", 200)
mycart.add_item("Orange", 150)

assert len(mycart.cart) == 3
assert mycart.cart['Papaya'] == 100
assert mycart.cart['Guava'] == 200
assert mycart.cart['Orange'] == 150

assert mycart.get_total_qty() == 450
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-83ddc5017d3d3393, points: 0, grade: true, solution: false -->
```python
# Test cases
mycart = ShoppingCart()

mycart.add_item("Papaya", 100)
mycart.add_item("Guava", 200)
mycart.add_item("Orange", 150)

mycart.remove_item("Orange",20)
mycart.remove_item("Papaya",200)

assert len(mycart.cart) == 2
assert mycart.cart['Guava'] == 200
assert mycart.cart['Orange'] == 130

assert mycart.get_total_qty() == 330
```

```python
# Test cases
mycart = ShoppingCart()

mycart.add_item("Papaya", 100)
mycart.add_item("Guava", 200)
mycart.add_item("Orange", 150)

mycart.remove_item("Orange",20)
mycart.remove_item("Papaya",200)

assert len(mycart.cart) == 2
assert mycart.cart['Guava'] == 200
assert mycart.cart['Orange'] == 130

assert mycart.get_total_qty() == 330
print('ok')
```

Output:
```
ok
```

**Q3:** Assume you want to develop a small banking system.

Given the `Account` class, write the `Bank` class in the provided section with the following methods:

- `__init__(profiles: dict)`: Constructs bank accounts based on the provided dictionary profiles. The keys of the accounts are the account names, which are keys in the profiles dictionary. The values are instances of the `Account` class with an initial balance (initial_balance) equal to the corresponding value in the profiles for that account name. 

For example,
If profiles is `{'Carol': 1000, 'Marta': -100}`, then the dictionary `accounts` in the Bank class would be `{'Carol': Account object with initial_balance 1000, 'Marta': Account object with initial_balance -100}`.

- `transfer(from_who: str, to_who: str, amt: float)`: Transfers money from the account named `from_who` to the account named `to_who` with the specified amount `amt`.

- get_names_in_debt(self): Returns a list of names that owe money to the bank, i.e., the account balances that are negative.

- `update_with_interest()`: Function for updating all accounts in the bank by calculating interest on both deposit and loan amounts.
python


![bank-accounts.png](attachment:bank-accounts.png)

```python
class Account:
  def __init__(self, initial_balance=0):
    self.__balance = initial_balance # private attribute starts with __

  def deposit(self, amount):
    self.__balance += amount

  def withdraw(self, amount):
    self.__balance -= amount        

  def get_balance(self):
    return self.__balance
```

<!-- grade_id: cell-7447a28c65bb37e1, points: 0, grade: false, solution: true -->
```python
class Bank():
    def __init__(self, profiles):
        self.accounts = {}
        if profiles is not {}:
            for key, value in profiles.items():
                self.accounts[key] = Account(value)
                
    def transfer(self, from_who, to_who, amt):
        # Sending money
        self.accounts[from_who].withdraw(amt)
        # Receiveing money
        self.accounts[to_who].deposit(amt)
        
    def get_names_in_debt(self):
        debt_names = []
        for key, account in self.accounts.items():
            if account.get_balance() < 0:
                debt_names.append(key)
        return debt_names
    
    def update_with_interest(self, interest_rate):
        for key, account in self.accounts.items():
            interest = account.get_balance() * interest_rate / 100
            account.deposit(interest)
```

<!-- grade_id: cell-0d272588c82ac692, points: 0, grade: true, solution: false -->
```python
# Test cases
csBank = Bank({'Carol':1000,'Marta':-100})
assert csBank.accounts['Carol'].get_balance() == 1000
assert csBank.accounts['Marta'].get_balance() == -100
```

<!-- grade_id: cell-c25ba2eec06a1847, points: 0, grade: true, solution: false -->
```python
# Test cases
csBank = Bank({'Carol':1000,'Marta':-100})
csBank.transfer('Carol','Marta',500)
assert csBank.accounts['Carol'].get_balance() == 500
assert csBank.accounts['Marta'].get_balance() == 400
```

<!-- grade_id: cell-9996078df9df63e5, points: 0, grade: true, solution: false -->
```python
# Test cases
csBank = Bank({'Carol':1000,'Marta':-100})
assert csBank.get_names_in_debt() == ['Marta']
```

<!-- grade_id: cell-cd6214f2135d069a, points: 0, grade: true, solution: false -->
```python
# Test cases
csBank = Bank({'Carol':1000,'Marta':-500})
csBank.update_with_interest(1.5)
assert csBank.accounts['Carol'].get_balance() == 1015
assert csBank.accounts['Marta'].get_balance() == -507.5,  csBank.accounts['Marta'].get_balance()
```

**Q2:** Imagine you are building a simple library management system. Implement three classes: `Book`, `Author`, and `Library`.

The Book class should have the following attributes:
- `title (str)`: the title of the book
- `isbn (str)`: the International Standard Book Number of the book
- `author (Author)`: the author of the book (an instance of the Author class)
- `available_copies (int)`: the number of available copies of the book

It should also have the following methods:
- `display_info()`: prints out the title, ISBN, author's name, and the number of available copies of the book.
- `borrow_copy()`: Decreases the number of available copies by 1 when a book is borrowed. If no copies are available, print a message indicating that the book is currently unavailable.
- `return_copy()`: Increases the number of available copies by 1 when a book is returned.


The Author class should have the following attributes:
- `name (str)`: the name of the author
- `birth_year (int)`: the birth year of the author

It should also have a method:
- `display_info()`: prints out the author's name and birth year.

The Library class should have the following attributes:
- `books (list)`: a list containing instances of the Book class

It should also have the following methods:
- `add_book(book: Book)`: adds a book to the library's collection
- `display_books()`: displays information about all the books in the library
- `search_books_by_author(author_name: str)`: Takes an author's name as input and prints the titles of all books in the library written by that author.
- `search_books_by_title(title: str)`: Takes a title as input and prints information about the book with that title in the library.
- `total_available_copies()`: Calculates and returns the total number of available copies for all books in the library.

Implement the classes and methods described above and create an example demonstrating the use of class composition with these classes.


<!-- grade_id: cell-5dd46a2d9399a4bb, points: 0, grade: false, solution: true -->
```python
class Book:
    def __init__(self, title, isbn, author, available_copies):
        self.title = title
        self.isbn = isbn
        self.author = author
        self.available_copies = available_copies
    
    def display_info(self):
        return f'{self.title}, {self.isbn}, {self.author.name}, {self.available_copies}'
    
    def borrow_copy(self):
        if self.available_copies > 0:
            self.available_copies -= 1
            return 'borrowed successfully'
        return 'currently unavailable'
    
    def return_copy(self):
        self.available_copies += 1
        return 'returned successfully!'

class Author:
    def __init__(self, name, birth_year):
        self.name = name
        self.birth_year = birth_year
        
    def display_info(self):
        return f'Author: {self.name}, Year: {self.birth_year}'
    
class Library:
    def __init__(self, books = []):
        self.books = {}
        if len(books) > 0:
            for book in books:
                self.books[book.isbn] = book
    
    def add_book(self, book):
        self.books[book.isbn] = book
        
    def display_books(self):
        books_information = ''
        for book in self.books.values():
            book_information = f'{book.display_info()}\n'
            books_information += book_information
        return books_information
    
    def search_books_by_author(self, author_name):
        books_title = ''
        for book in self.books.values():
            if book.author.name == author_name:
                books_title += f'{book.title}\n'
        if books_title != '':
            return books_title
        return 'No books found'
    
    def search_books_by_title(self, title):
        for book in self.books.values():
            if book.title == title:
                return book.display_info()
        return 'No book found'
                
    def total_available_copies(self):
        total_copies = 0
        for book in self.books.values():
            total_copies += book.available_copies
        return total_copies
```

```python
# Example Usage
author1 = Author("John Doe", 1980)
book1 = Book("Introduction to Python", "978-1-234567-89-0", author1, 10)
book2 = Book("Data Structures in Python", "978-1-234567-89-1", author1, 5)
book3 = Book("Advanced Python Programming", "978-1-234567-89-2", author1, 1)


library = Library()
library.add_book(book1)
library.add_book(book2)
library.add_book(book3)

# Display information about the books in the library
print(library.display_books())

# Borrow a copy of a book
print(library.books[book1.isbn].borrow_copy())

# Return a copy of a book
print(library.books[book2.isbn].return_copy())

# Display updated information about the books in the library
print(library.display_books())

# Search for books by a specific author
print(library.search_books_by_author("John Doe"))

# Search for a book by title
print(library.search_books_by_title("Introduction to Python"))

# Get the total number of available copies in the library
print(library.total_available_copies())

```

Output:
```
Introduction to Python, 978-1-234567-89-0, John Doe, 10
Data Structures in Python, 978-1-234567-89-1, John Doe, 5
Advanced Python Programming, 978-1-234567-89-2, John Doe, 1

borrowed successfully
returned successfully!
Introduction to Python, 978-1-234567-89-0, John Doe, 9
Data Structures in Python, 978-1-234567-89-1, John Doe, 6
Advanced Python Programming, 978-1-234567-89-2, John Doe, 1

Introduction to Python
Data Structures in Python
Advanced Python Programming

Introduction to Python, 978-1-234567-89-0, John Doe, 9
16
```

```python
# Example Usage
author1 = Author("John Doe", 1980)
book1 = Book("Introduction to Python", "978-1-234567-89-0", author1, 10)
book2 = Book("Data Structures in Python", "978-1-234567-89-1", author1, 5)
book3 = Book("Advanced Python Programming", "978-1-234567-89-2", author1, 1)


library = Library()
library.add_book(book1)
library.add_book(book2)
library.add_book(book3)

# Display information about the books in the library
print('Display information about the books in the library')
print(library.display_books())

# Borrow a copy of a book
print('Borrow a copy of a book')
print(library.books[book1.isbn].borrow_copy())

# Return a copy of a book
print('Return a copy of a book')
print(library.books[book2.isbn].return_copy())

# Display updated information about the books in the library
print('Display updated information about the books in the library')
print(library.display_books())

# Search for books by a specific author
print('Search for books by a specific author')
print(library.search_books_by_author("John Doe"))

# Search for a book by title
print('Search for a book by title')
print(library.search_books_by_title("Introduction to Python"))

# Get the total number of available copies in the library
print('Get the total number of available copies in the library')
print(library.total_available_copies())

```

Output:
```
Display information about the books in the library
Introduction to Python, 978-1-234567-89-0, John Doe, 10
Data Structures in Python, 978-1-234567-89-1, John Doe, 5
Advanced Python Programming, 978-1-234567-89-2, John Doe, 1

Borrow a copy of a book
borrowed successfully
Return a copy of a book
returned successfully!
Display updated information about the books in the library
Introduction to Python, 978-1-234567-89-0, John Doe, 9
Data Structures in Python, 978-1-234567-89-1, John Doe, 6
Advanced Python Programming, 978-1-234567-89-2, John Doe, 1

Search for books by a specific author
Introduction to Python
Data Structures in Python
Advanced Python Programming

Search for a book by title
Introduction to Python, 978-1-234567-89-0, John Doe, 9
Get the total number of available copies in the library
16
```

In the context of Python's unittest framework, each test case is executed in a new instance of the test case class. When a test case is run, the `setUp` method is called before each test method, and the `tearDown` method is called after each test method. 

Test cases were created as follows.

**Author Class Test Cases:**
- Test Case 1.1: Create an Author instance with a valid name and birth year.
- Test Case 1.2: Verify that display_info() method returns a string.

**Book Class Test Cases:**
- Test Case 2.1: Create a Book instance with a valid title, ISBN, Author instance, and available copies.
- Test Case 2.2: Verify that display_info() method returns a string.
- Test Case 2.3: Borrow a copy of a book with available copies.
- Test Case 2.4: Return a copy of a book.
- Test Case 2.5: Attempt to borrow a copy of a book with no available copies.

**Library Class Test Cases:**
- Test Case 3.1: Create a Library instance.
- Test Case 3.2: Add a Book instance to the Library.
- Test Case 3.3: Verify that display_books() method returns a string.
- Test Case 3.4: Search for books by a valid author's name.
- Test Case 3.5: Search for a book by a valid title.
- Test Case 3.6: Attempt to search for books by an invalid author's name.
- Test Case 3.7: Attempt to search for a non-existing book by title.
- Test Case 3.8: Verify that total_available_copies() method returns an integer.

```python
import unittest

class TestAuthor(unittest.TestCase):
    
    def setUp(self): 
        self.author = Author("John Doe", 1980)

    def test_author_creation(self):
        # Test Case 1.1
        self.assertIsInstance(self.author, Author)
        print("Test Case 1.1 passed")

    def test_author_display_info(self):
        # Test Case 1.2
        author_info = self.author.display_info()
        self.assertIsInstance(author_info, str)
        self.assertIn("John Doe", author_info)
        self.assertIn("1980", author_info)
        print("Test Case 1.2 passed")

class TestBook(unittest.TestCase):    
  
    def setUp(self): 
        self.author = Author("John Doe", 1980)
        self.book1 = Book("Introduction to Python", "978-1-234567-89-0", self.author, 10)
        self.book2 = Book("Data Structures in Python", "978-1-234567-89-1", self.author, 5)
        self.book3 = Book("Advanced Python Programming", "978-1-234567-89-2", self.author, 1)
        
    def test_book_creation(self):
        # Test Case 2.1
        self.assertIsInstance(self.book1, Book)
        print("Test Case 2.1 passed")

    def test_book_display_info(self):
        # Test Case 2.2
        book_info = self.book1.display_info()
        self.assertIsInstance(book_info, str)
        self.assertIn("Introduction to Python", book_info)
        self.assertIn("978-1-234567-89-0", book_info)
        self.assertIn("John Doe", book_info)
        self.assertIn("10", book_info)
        print("Test Case 2.2 passed")

    def test_borrow_copy(self):
        # Test Case 2.3
        borrow_result = self.book1.borrow_copy()
        self.assertIsInstance(borrow_result, str)
        self.assertIn("borrowed successfully", borrow_result)
        print("Test Case 2.3 passed")

    def test_return_copy(self):
        # Test Case 2.4
        return_result = self.book2.return_copy()
        self.assertIsInstance(return_result, str)
        self.assertIn("returned successfully", return_result)
        print("Test Case 2.4 passed")

    def test_borrow_copy_no_available_copies(self):
        # Test Case 2.5
        self.book3.borrow_copy()  # Borrow the only available copy
        borrow_result = self.book3.borrow_copy()
        self.assertIsInstance(borrow_result, str)
        self.assertIn("currently unavailable", borrow_result)
        print("Test Case 2.5 passed")

class TestLibrary(unittest.TestCase):
  
    def setUp(self): 
        self.author = Author("John Doe", 1980)
        self.book1 = Book("Introduction to Python", "978-1-234567-89-0", self.author, 10)
        self.book2 = Book("Data Structures in Python", "978-1-234567-89-1", self.author, 5)
        self.book3 = Book("Advanced Python Programming", "978-1-234567-89-2", self.author, 1)
        self.library = Library()
        self.library.add_book(self.book1)
        self.library.add_book(self.book2)
        self.library.add_book(self.book3)
        
    def test_library_creation(self):
        # Test Case 3.1
        self.assertIsInstance(self.library, Library)
        print("Test Case 3.1 passed")

    def test_add_book(self):
        # Test Case 3.2
        self.assertEqual(len(self.library.books), 3)
        print("Test Case 3.2 passed")

    def test_display_books(self):
        # Test Case 3.3
        library_info = self.library.display_books()
        self.assertIsInstance(library_info, str)
        self.assertIn("Introduction to Python", library_info)
        self.assertIn("Data Structures in Python", library_info)
        self.assertIn("Advanced Python Programming", library_info)
        print("Test Case 3.3 passed")

    def test_search_books_by_author(self):
        # Test Case 3.4
        author_books = self.library.search_books_by_author("John Doe")
        self.assertIsInstance(author_books, str)
        self.assertIn("Introduction to Python", author_books)
        print("Test Case 3.4 passed")

    def test_search_books_by_title(self):
        # Test Case 3.5
        book_info = self.library.search_books_by_title("Introduction to Python")
        self.assertIsInstance(book_info, str)
        self.assertIn("Introduction to Python", book_info)
        print("Test Case 3.5 passed")

    def test_search_books_by_invalid_author(self):
        # Test Case 3.6
        not_found_msg = self.library.search_books_by_author("Unknown Author")
        self.assertIsInstance(not_found_msg, str)
        self.assertIn("No books found", not_found_msg)
        print("Test Case 3.6 passed")

    def test_search_books_by_non_existing_title(self):
        # Test Case 3.7
        not_found_msg = self.library.search_books_by_title("Non-Existing Title")
        self.assertIsInstance(not_found_msg, str)
        self.assertIn("No book found", not_found_msg)
        print("Test Case 3.7 passed")

    def test_total_available_copies(self):
        # Test Case 3.8
        isbn = "978-1-234567-89-0"
        self.library.books[isbn].borrow_copy()
        self.library.books[isbn].borrow_copy()
        self.library.books[isbn].return_copy()
        total_copies = self.library.total_available_copies()
        self.assertIsInstance(total_copies, int)
        self.assertEqual(total_copies, 15)
        print("Test Case 3.8 passed")
```

<!-- grade_id: cell-71892696ad6a72a0, points: 1, grade: true, solution: false -->
```python
# Load and run TestAuthor
result = unittest.TextTestRunner().run(unittest.TestLoader().loadTestsFromTestCase(TestAuthor))
```

Output:
```
..
----------------------------------------------------------------------
Ran 2 tests in 0.004s

OK
Test Case 1.1 passed
Test Case 1.2 passed
```

<!-- grade_id: cell-9ef6ef2a9fa48788, points: 2, grade: true, solution: false -->
```python
# Load and run TestBook
unittest.TextTestRunner().run(unittest.TestLoader().loadTestsFromTestCase(TestBook))
```

Output:
```
.....
----------------------------------------------------------------------
Ran 5 tests in 0.008s

OK
Test Case 2.1 passed
Test Case 2.2 passed
Test Case 2.3 passed
Test Case 2.5 passed
Test Case 2.4 passed
<unittest.runner.TextTestResult run=5 errors=0 failures=0>
```

<!-- grade_id: cell-ee24ead1bc0f5132, points: 2, grade: true, solution: false -->
```python
# Load and run TestLibrary
unittest.TextTestRunner().run(unittest.TestLoader().loadTestsFromTestCase(TestLibrary))
```

Output:
```
........
----------------------------------------------------------------------
Ran 8 tests in 0.007s

OK
Test Case 3.2 passed
Test Case 3.3 passed
Test Case 3.1 passed
Test Case 3.4 passed
Test Case 3.6 passed
Test Case 3.7 passed
Test Case 3.5 passed
Test Case 3.8 passed
<unittest.runner.TextTestResult run=8 errors=0 failures=0>
```

```python

```

```python

```